Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` (you should of course delete `raise NotImplementedError()` which is there only as reminder), while not modifying the other cells (but you should run them to check the output you obtain). Please also fill in your group number, your full names, and your VU IDs below:

```python
STUDENT_FIRST_NAME = "Thu Trang" # e.g. "John" (no "J", no "J.", no "John S.M.")
STUDENT_LAST_NAME = "Luu"  # e.g. "Smith"
VU_ID_NUM = "2695303"          # e.g. "2789012"
```

---

# Assignment 0

The goal of this mock assignment is just to make you familiar with the rather rigid structure in which you need to submit the three assignments of this course. 

The fact that is a **mock** assignment means that it *does not count towards your final grade*. Nonetheless, you are warmly encourage to quickly write your solution and submit it so that you can avoid undesirable submission issues or autograder issues for the next three assignments.

We assume that the folder that you work in is the one obtained by unzipping the given ``assignment0.zip`` file and thus has the following structure.

<code>
assignment0.ipynb
data/a0_input0.txt
data/a0_input1.csv
</code>

In particular, the subfolder data should have ``2`` files.

## Important remarks
- **Working together**: You are meant to work individually on the first three assignments. You can, of course, brainstorm ideas and discuss issues with your fellow students, but you are required to write your solutions individually. 
- **Plagiarism**: All your code will be automatically scanned for plagiarism. Furthermore, using the internet as a passive resource is allowed. This means that you can search for help there and partially copy code, as long as you explicitly acknowledge inside your Jupyter notebook which parts have been copied and from where. 
- **Performance**: You should optimize the computational performance of your functions. Specifically, when grading the assignments, we set a hard limit of 5 seconds for each cell execution. This should be sufficient to cover each of the test cases. Function calls that take longer than 5 seconds will not be awarded any points.
- **Code styling**: Your implementation will not be checked for style. However, we do encourage you to practice good code styling. See, for example, https://docs.python-guide.org/writing/style/.
- **Chronological run**: All outputs should be repeatable by doing one full “chronological” run of the notebook without any manual changes to code blocks, including parameters. (try it yourself by clicking ``Kernel -> Restart and run all``, which should give the result as handed in).
- **Handing in**: Hand in the .ipynb file of your notebook on ``Canvas`` before the assignment deadline. 
- **Other questions**: If you have doubts/questions about the assignments, feel free to ask them in [this discussion thread](https://canvas.vu.nl/courses/60060/discussion_topics/531826) so that everyone can will be able to see them and our answers. 

```python
# Remember to run this cell before anything else

from numpy.testing import assert_equal
```

## Exercise 1: Reading a file

Given a filename, write a function that returns the content of the corresponding file.

<!-- grade_id: Q1-code, points: 0, grade: false, solution: true -->
```python
def read_file (filename):
    with open (filename) as f:
        f_list = f.read()
    
    return f_list
```

<!-- grade_id: Q1-test, points: 1, grade: true, solution: false -->
```python
# Test your function on a few possible inputs

assert_equal(read_file("data/a0_input0.txt"),"Welcome to the Project Big Data course!")
assert_equal(read_file("data/a0_input1.csv"),"1,2,3\n4,5,6\n7,8,9")

```

```python
def test (var):
    if type(var) == str:
        print('type string')
    else:
        print('type not string')
```

```python
test ("hello")
```

Output:
```
type string
```

## Exercise 2: Timing your function

Write a Python function that return a list of all even numbers between $0$ (included) and $n$ (included, if $n$ is even). Test your function for large integer values $n$ using the magic word ``%time even_numbers(n)``.

Your function will be tested for $n$ up to $10^7$, so make sure it is efficient and runs fast enough (< 5 seconds).

<!-- grade_id: Q2-code, points: 0, grade: false, solution: true -->
```python
def even_numbers(n):    
    # create input list
    input = list(range(n+1))
    
    # find evens
    result = input[::2]
    
    return result
```

<!-- grade_id: Q2-test, points: 1, grade: true, solution: false -->
```python
# Test the correctness of your function on a few possible inputs while 

assert_equal(even_numbers(10), [0, 2, 4, 6, 8, 10])
assert_equal(even_numbers(40), [0, 2, 4, 6, 8, 10, 12, 14, 16, 18, 20, 22, 24, 26, 28, 30, 32, 34, 36, 38, 40])

# To check how long your function takes you can use the following

%time even_numbers(40)

# Check what the running time is for larger inputs up to 10^7 and make sure is less than 2 seconds. 
# Note that the correctedness of the output is not checked.

from time import perf_counter

lengths = [10, 100, 1000, 10000, 100000, 1000000, 10000000]

for i, n in enumerate(lengths):
    
    t_start = perf_counter()
    e = even_numbers(n)
    t_stop = perf_counter()
    print(f"\nTest {i}. Your function took {t_stop-t_start} seconds when given n={n} as input.")

```

Output:
```
Wall time: 0 ns

Test 0. Your function took 2.000000000279556e-06 seconds when given n=10 as input.

Test 1. Your function took 2.299999999344493e-06 seconds when given n=100 as input.

Test 2. Your function took 5.489999999941375e-05 seconds when given n=1000 as input.

Test 3. Your function took 0.00028149999999982356 seconds when given n=10000 as input.

Test 4. Your function took 0.010506399999999694 seconds when given n=100000 as input.

Test 5. Your function took 0.029455200000000126 seconds when given n=1000000 as input.

Test 6. Your function took 0.4136793000000001 seconds when given n=10000000 as input.
```